fof(pel55_11,axiom,
    agatha != butler,
    file('/Users/schulz/SOURCES/Projects/PyRes/EXAMPLES/PUZ001+1.p',pel55_11) ).

cnf(c7,axiom,
    agatha != butler,
    inference(split_conjunct,[status(thm)],[pel55_11]) ).

fof(pel55_7,axiom,
    ! [X] :
      ( X != butler
     => hates(agatha,X) ),
    file('/Users/schulz/SOURCES/Projects/PyRes/EXAMPLES/PUZ001+1.p',pel55_7) ).

fof(c19,plain,
    ! [X] :
      ( X = butler
      | hates(agatha,X) ),
    inference(fof_nnf,[status(thm)],[pel55_7]) ).

fof(c20,plain,
    ! [X6] :
      ( X6 = butler
      | hates(agatha,X6) ),
    inference(variable_rename,[status(thm)],[c19]) ).

cnf(c21,plain,
    ( X38 = butler
    | hates(agatha,X38) ),
    inference(split_conjunct,[status(thm)],[c20]) ).

cnf(c53,plain,
    hates(agatha,agatha),
    inference(resolution,[status(thm)],[c21,c7]) ).

fof(pel55_6,axiom,
    ! [X] :
      ( hates(agatha,X)
     => ~ hates(charles,X) ),
    file('/Users/schulz/SOURCES/Projects/PyRes/EXAMPLES/PUZ001+1.p',pel55_6) ).

fof(c22,plain,
    ! [X] :
      ( hates(agatha,X)
     => ~ hates(charles,X) ),
    inference(fof_simplification,[status(thm)],[pel55_6]) ).

fof(c23,plain,
    ! [X] :
      ( ~ hates(agatha,X)
      | ~ hates(charles,X) ),
    inference(fof_nnf,[status(thm)],[c22]) ).

fof(c24,plain,
    ! [X7] :
      ( ~ hates(agatha,X7)
      | ~ hates(charles,X7) ),
    inference(variable_rename,[status(thm)],[c23]) ).

cnf(c25,plain,
    ( ~ hates(agatha,X43)
    | ~ hates(charles,X43) ),
    inference(split_conjunct,[status(thm)],[c24]) ).

cnf(reflexivity,axiom,
    X14 = X14,
    theory(equality) ).

fof(pel55_1,axiom,
    ? [X] :
      ( lives(X)
      & killed(X,agatha) ),
    file('/Users/schulz/SOURCES/Projects/PyRes/EXAMPLES/PUZ001+1.p',pel55_1) ).

fof(c39,plain,
    ? [X13] :
      ( lives(X13)
      & killed(X13,agatha) ),
    inference(variable_rename,[status(thm)],[pel55_1]) ).

fof(c40,plain,
    ( lives(skolem0002)
    & killed(skolem0002,agatha) ),
    inference(skolemize,[status(esa)],[c39]) ).

cnf(c42,plain,
    killed(skolem0002,agatha),
    inference(split_conjunct,[status(thm)],[c40]) ).

fof(pel55_4,axiom,
    ! [X,Y] :
      ( killed(X,Y)
     => hates(X,Y) ),
    file('/Users/schulz/SOURCES/Projects/PyRes/EXAMPLES/PUZ001+1.p',pel55_4) ).

fof(c30,plain,
    ! [X,Y] :
      ( ~ killed(X,Y)
      | hates(X,Y) ),
    inference(fof_nnf,[status(thm)],[pel55_4]) ).

fof(c31,plain,
    ! [X10,X11] :
      ( ~ killed(X10,X11)
      | hates(X10,X11) ),
    inference(variable_rename,[status(thm)],[c30]) ).

cnf(c32,plain,
    ( ~ killed(X25,X24)
    | hates(X25,X24) ),
    inference(split_conjunct,[status(thm)],[c31]) ).

cnf(c46,plain,
    hates(skolem0002,agatha),
    inference(resolution,[status(thm)],[c32,c42]) ).

cnf(c2,axiom,
    ( X39 != X40
    | X41 != X42
    | ~ hates(X39,X41)
    | hates(X40,X42) ),
    theory(equality) ).

cnf(c67,plain,
    ( skolem0002 != X72
    | agatha != X73
    | hates(X72,X73) ),
    inference(resolution,[status(thm)],[c2,c46]) ).

cnf(c155,plain,
    ( skolem0002 != X83
    | hates(X83,agatha) ),
    inference(resolution,[status(thm)],[c67,reflexivity]) ).

fof(pel55_5,axiom,
    ! [X,Y] :
      ( killed(X,Y)
     => ~ richer(X,Y) ),
    file('/Users/schulz/SOURCES/Projects/PyRes/EXAMPLES/PUZ001+1.p',pel55_5) ).

fof(c26,plain,
    ! [X,Y] :
      ( killed(X,Y)
     => ~ richer(X,Y) ),
    inference(fof_simplification,[status(thm)],[pel55_5]) ).

fof(c27,plain,
    ! [X,Y] :
      ( ~ killed(X,Y)
      | ~ richer(X,Y) ),
    inference(fof_nnf,[status(thm)],[c26]) ).

fof(c28,plain,
    ! [X8,X9] :
      ( ~ killed(X8,X9)
      | ~ richer(X8,X9) ),
    inference(variable_rename,[status(thm)],[c27]) ).

cnf(c29,plain,
    ( ~ killed(X23,X22)
    | ~ richer(X23,X22) ),
    inference(split_conjunct,[status(thm)],[c28]) ).

fof(pel55_10,axiom,
    ! [X] :
    ? [Y] : ~ hates(X,Y),
    file('/Users/schulz/SOURCES/Projects/PyRes/EXAMPLES/PUZ001+1.p',pel55_10) ).

fof(c8,plain,
    ! [X] :
    ? [Y] : ~ hates(X,Y),
    inference(fof_simplification,[status(thm)],[pel55_10]) ).

fof(c9,plain,
    ! [X2] :
    ? [X3] : ~ hates(X2,X3),
    inference(variable_rename,[status(thm)],[c8]) ).

fof(c10,plain,
    ! [X2] : ~ hates(X2,skolem0001(X2)),
    inference(skolemize,[status(esa)],[c9]) ).

cnf(c11,plain,
    ~ hates(X18,skolem0001(X18)),
    inference(split_conjunct,[status(thm)],[c10]) ).

fof(pel55_9,axiom,
    ! [X] :
      ( hates(agatha,X)
     => hates(butler,X) ),
    file('/Users/schulz/SOURCES/Projects/PyRes/EXAMPLES/PUZ001+1.p',pel55_9) ).

fof(c12,plain,
    ! [X] :
      ( ~ hates(agatha,X)
      | hates(butler,X) ),
    inference(fof_nnf,[status(thm)],[pel55_9]) ).

fof(c13,plain,
    ! [X4] :
      ( ~ hates(agatha,X4)
      | hates(butler,X4) ),
    inference(variable_rename,[status(thm)],[c12]) ).

cnf(c14,plain,
    ( ~ hates(agatha,X32)
    | hates(butler,X32) ),
    inference(split_conjunct,[status(thm)],[c13]) ).

cnf(c57,plain,
    ( X50 = butler
    | hates(butler,X50) ),
    inference(resolution,[status(thm)],[c21,c14]) ).

cnf(c85,plain,
    skolem0001(butler) = butler,
    inference(resolution,[status(thm)],[c57,c11]) ).

fof(pel55_8,axiom,
    ! [X] :
      ( ~ richer(X,agatha)
     => hates(butler,X) ),
    file('/Users/schulz/SOURCES/Projects/PyRes/EXAMPLES/PUZ001+1.p',pel55_8) ).

fof(c15,plain,
    ! [X] :
      ( ~ richer(X,agatha)
     => hates(butler,X) ),
    inference(fof_simplification,[status(thm)],[pel55_8]) ).

fof(c16,plain,
    ! [X] :
      ( richer(X,agatha)
      | hates(butler,X) ),
    inference(fof_nnf,[status(thm)],[c15]) ).

fof(c17,plain,
    ! [X5] :
      ( richer(X5,agatha)
      | hates(butler,X5) ),
    inference(variable_rename,[status(thm)],[c16]) ).

cnf(c18,plain,
    ( richer(X33,agatha)
    | hates(butler,X33) ),
    inference(split_conjunct,[status(thm)],[c17]) ).

cnf(c49,plain,
    richer(skolem0001(butler),agatha),
    inference(resolution,[status(thm)],[c18,c11]) ).

cnf(c3,axiom,
    ( X46 != X47
    | X48 != X49
    | ~ richer(X46,X48)
    | richer(X47,X49) ),
    theory(equality) ).

cnf(c79,plain,
    ( skolem0001(butler) != X107
    | agatha != X106
    | richer(X107,X106) ),
    inference(resolution,[status(thm)],[c3,c49]) ).

cnf(c273,plain,
    ( agatha != X108
    | richer(butler,X108) ),
    inference(resolution,[status(thm)],[c79,c85]) ).

cnf(c275,plain,
    richer(butler,agatha),
    inference(resolution,[status(thm)],[c273,reflexivity]) ).

cnf(c278,plain,
    ~ killed(butler,agatha),
    inference(resolution,[status(thm)],[c275,c29]) ).

cnf(c1,axiom,
    ( X34 != X35
    | X36 != X37
    | ~ killed(X34,X36)
    | killed(X35,X37) ),
    theory(equality) ).

cnf(c51,plain,
    ( skolem0002 != X54
    | agatha != X53
    | killed(X54,X53) ),
    inference(resolution,[status(thm)],[c1,c42]) ).

cnf(c103,plain,
    ( skolem0002 != X60
    | killed(X60,agatha) ),
    inference(resolution,[status(thm)],[c51,reflexivity]) ).

fof(pel55,conjecture,
    killed(agatha,agatha),
    file('/Users/schulz/SOURCES/Projects/PyRes/EXAMPLES/PUZ001+1.p',pel55) ).

fof(c4,negated_conjecture,
    ~ killed(agatha,agatha),
    inference(assume_negation,[status(cth)],[pel55]) ).

fof(c5,negated_conjecture,
    ~ killed(agatha,agatha),
    inference(fof_simplification,[status(thm)],[c4]) ).

cnf(c6,negated_conjecture,
    ~ killed(agatha,agatha),
    inference(split_conjunct,[status(thm)],[c5]) ).

cnf(c41,plain,
    lives(skolem0002),
    inference(split_conjunct,[status(thm)],[c40]) ).

fof(pel55_3,axiom,
    ! [X] :
      ( lives(X)
     => ( X = agatha
        | X = butler
        | X = charles ) ),
    file('/Users/schulz/SOURCES/Projects/PyRes/EXAMPLES/PUZ001+1.p',pel55_3) ).

fof(c33,plain,
    ! [X] :
      ( ~ lives(X)
      | X = agatha
      | X = butler
      | X = charles ),
    inference(fof_nnf,[status(thm)],[pel55_3]) ).

fof(c34,plain,
    ! [X12] :
      ( ~ lives(X12)
      | X12 = agatha
      | X12 = butler
      | X12 = charles ),
    inference(variable_rename,[status(thm)],[c33]) ).

cnf(c35,plain,
    ( ~ lives(X51)
    | X51 = agatha
    | X51 = butler
    | X51 = charles ),
    inference(split_conjunct,[status(thm)],[c34]) ).

cnf(c94,plain,
    ( skolem0002 = agatha
    | skolem0002 = butler
    | skolem0002 = charles ),
    inference(resolution,[status(thm)],[c35,c41]) ).

cnf(c365,plain,
    ( skolem0002 = butler
    | skolem0002 = charles
    | killed(agatha,agatha) ),
    inference(resolution,[status(thm)],[c94,c103]) ).

cnf(c4177,plain,
    ( skolem0002 = butler
    | skolem0002 = charles ),
    inference(resolution,[status(thm)],[c365,c6]) ).

cnf(c4201,plain,
    ( skolem0002 = charles
    | killed(butler,agatha) ),
    inference(resolution,[status(thm)],[c4177,c103]) ).

cnf(c4308,plain,
    skolem0002 = charles,
    inference(resolution,[status(thm)],[c4201,c278]) ).

cnf(c4334,plain,
    hates(charles,agatha),
    inference(resolution,[status(thm)],[c4308,c155]) ).

cnf(c4361,plain,
    ~ hates(agatha,agatha),
    inference(resolution,[status(thm)],[c4334,c25]) ).

cnf(c4408,plain,
    $false,
    inference(resolution,[status(thm)],[c4361,c53]) ).
